% -------------------------------------------------------------------------
% Converged at ZigZag number 2:1 with similarity 0.95
% The original and final NL have similarity: 0.95
% NL2L model:       OpenAI---gpt-5-chat-latest
% L2NL model:       OpenAI---gpt-5-chat-latest
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
John was in noncompliance with the bank overdraft policy.
The charging of a fee led to a deposit.
As a result of the deposit, the bank reduced the penalty.
What led to a deposit?

% -------------------------------------------------------------------------
% The final NL is :
John is in noncompliance with the overdraft policy.  
Charging a fee led to a deposit.  
A deposit resulted in a reduction of a penalty, and the bank reduced the penalty.  
Is there an event that led to a deposit?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both segments describe the same sequence of events: John’s noncompliance with the overdraft policy, a fee being charged that led to a deposit, and the bank reducing a penalty as a result. The only differences are minor wording and tense changes (“was” vs. “is,” “as a result of” vs. “resulted in”), and the second text adds a slightly rephrased question. These do not significantly alter the meaning.
% -------------------------------------------------------------------------
% The final logic is :
tff(t_type,type,
    t: $tType ).

tff(event_type,type,
    event: $tType ).

tff(agent_type,type,
    agent: $tType ).

tff(policy_type,type,
    policy: $tType ).

tff(john_decl,type,
    john: agent ).

tff(bank_decl,type,
    bank: agent ).

tff(overdraft_policy_decl,type,
    overdraft_policy: policy ).

tff(noncompliance_decl,type,
    noncompliance: event ).

tff(charging_fee_decl,type,
    charging_fee: event ).

tff(deposit_decl,type,
    deposit: event ).

tff(reduce_penalty_decl,type,
    reduce_penalty: event ).

tff(in_noncompliance_with,type,
    in_noncompliance_with: ( agent * policy ) > $o ).

tff(led_to,type,
    led_to: ( event * event ) > $o ).

tff(as_result_of,type,
    as_result_of: ( event * event ) > $o ).

tff(reduced,type,
    reduced: ( agent * event ) > $o ).

tff(s1,axiom,
    in_noncompliance_with(john,overdraft_policy) ).

tff(s2,axiom,
    led_to(charging_fee,deposit) ).

tff(s3,axiom,
    ( as_result_of(deposit,reduce_penalty)
    & reduced(bank,reduce_penalty) ) ).

tff(q1,conjecture,
    ? [X: event] : led_to(X,deposit) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, t: $tType).
tff(type_def_6, type, event: $tType).
tff(type_def_7, type, agent: $tType).
tff(type_def_8, type, policy: $tType).
tff(func_def_0, type, john: agent).
tff(func_def_1, type, bank: agent).
tff(func_def_2, type, overdraft_policy: policy).
tff(func_def_3, type, noncompliance: event).
tff(func_def_4, type, charging_fee: event).
tff(func_def_5, type, deposit: event).
tff(func_def_6, type, reduce_penalty: event).
tff(pred_def_1, type, in_noncompliance_with: (agent * policy) > $o).
tff(pred_def_2, type, led_to: (event * event) > $o).
tff(pred_def_3, type, as_result_of: (event * event) > $o).
tff(pred_def_4, type, reduced: (agent * event) > $o).
tff(f12,plain,(
  $false),
  inference(subsumption_resolution,[],[f10,f11])).
tff(f11,plain,(
  ( ! [X0 : event] : (~led_to(X0,deposit)) )),
  inference(cnf_transformation,[],[f9])).
tff(f9,plain,(
  ! [X0 : event] : ~led_to(X0,deposit)),
  inference(ennf_transformation,[],[f5])).
tff(f5,negated_conjecture,(
  ~? [X0 : event] : led_to(X0,deposit)),
  inference(negated_conjecture,[status(cth)],[f4])).
tff(f4,conjecture,(
  ? [X0 : event] : led_to(X0,deposit)),
  file('/tmp/lxjfsNcsgW/SOT_wK4ykV',unknown)).
tff(f10,plain,(
  led_to(charging_fee,deposit)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  led_to(charging_fee,deposit)),
  file('/tmp/lxjfsNcsgW/SOT_wK4ykV',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-AP/02-Bank.THM.txt,True,0.95,2,True,0.95,1,THM,Unknown,0,0,0,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 3:2 with similarity 0.98
% The original and final NL have similarity: 0.93
% NL2L model:       OpenAI---gpt-5-chat-latest
% L2NL model:       OpenAI---gpt-5-chat-latest
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
John was in noncompliance with the bank overdraft policy.
The charging of a fee led to a deposit.
As a result of the deposit, the bank reduced the penalty.
What led to a deposit?

% -------------------------------------------------------------------------
% The final NL is :
John is in noncompliance with the overdraft policy.  
The fee charging led to the deposit.  
The deposit resulted in a penalty reduction, and the bank reduced the penalty reduction.  
Is there an event that led to the deposit?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both passages describe nearly the same sequence of events: John violates the overdraft policy, a fee is charged, that charge leads to a deposit, and the deposit causes the bank to reduce a penalty. The main differences are stylistic and minor wording changes (“is” vs. “was,” “the overdraft policy” vs. “the bank overdraft policy”). The second text adds a redundant phrase (“the bank reduced the penalty reduction”) and slightly rephrases the final question, but the underlying meaning remains almost identical.
% -------------------------------------------------------------------------
% The final logic is :
tff(type_event,type,
    event: $tType ).

tff(type_entity,type,
    entity: $tType ).

tff(type_person,type,
    person: $tType ).

tff(type_policy,type,
    policy: $tType ).

tff(pred_in_noncompliance,type,
    in_noncompliance: ( person * policy ) > $o ).

tff(pred_led_to,type,
    led_to: ( event * event ) > $o ).

tff(pred_resulted_in,type,
    resulted_in: ( event * event ) > $o ).

tff(pred_reduced,type,
    reduced: ( entity * event ) > $o ).

tff(const_john,type,
    john: person ).

tff(const_overdraft_policy,type,
    overdraft_policy: policy ).

tff(const_fee_charging,type,
    fee_charging: event ).

tff(const_deposit,type,
    deposit: event ).

tff(const_penalty_reduction,type,
    penalty_reduction: event ).

tff(const_bank,type,
    bank: entity ).

tff(ax1,axiom,
    in_noncompliance(john,overdraft_policy) ).

tff(ax2,axiom,
    led_to(fee_charging,deposit) ).

tff(ax3,axiom,
    ( resulted_in(deposit,penalty_reduction)
    & reduced(bank,penalty_reduction) ) ).

tff(conj1,conjecture,
    ? [E: event] : led_to(E,deposit) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, event: $tType).
tff(type_def_6, type, entity: $tType).
tff(type_def_7, type, person: $tType).
tff(type_def_8, type, policy: $tType).
tff(func_def_0, type, john: person).
tff(func_def_1, type, overdraft_policy: policy).
tff(func_def_2, type, fee_charging: event).
tff(func_def_3, type, deposit: event).
tff(func_def_4, type, penalty_reduction: event).
tff(func_def_5, type, bank: entity).
tff(pred_def_1, type, in_noncompliance: (person * policy) > $o).
tff(pred_def_2, type, led_to: (event * event) > $o).
tff(pred_def_3, type, resulted_in: (event * event) > $o).
tff(pred_def_4, type, reduced: (entity * event) > $o).
tff(f12,plain,(
  $false),
  inference(resolution,[],[f10,f11])).
tff(f11,plain,(
  ( ! [X0 : event] : (~led_to(X0,deposit)) )),
  inference(cnf_transformation,[],[f9])).
tff(f9,plain,(
  ! [X0 : event] : ~led_to(X0,deposit)),
  inference(ennf_transformation,[],[f5])).
tff(f5,negated_conjecture,(
  ~? [X0 : event] : led_to(X0,deposit)),
  inference(negated_conjecture,[status(cth)],[f4])).
tff(f4,conjecture,(
  ? [X0 : event] : led_to(X0,deposit)),
  file('/tmp/5ofQsYFtjE/SOT_AHuryc',unknown)).
tff(f10,plain,(
  led_to(fee_charging,deposit)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  led_to(fee_charging,deposit)),
  file('/tmp/5ofQsYFtjE/SOT_AHuryc',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-AP/02-Bank.THM.txt,False,0.93,3,True,0.98,2,THM,Unknown,0,1,0,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 3:2 with similarity 1.00
% The original and final NL have similarity: 0.93
% NL2L model:       OpenAI---gpt-5-chat-latest
% L2NL model:       OpenAI---gpt-5-chat-latest
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
John was in noncompliance with the bank overdraft policy.
The charging of a fee led to a deposit.
As a result of the deposit, the bank reduced the penalty.
What led to a deposit?

% -------------------------------------------------------------------------
% The final NL is :
John was in noncompliance with the overdraft policy.  
The fee charging led to the deposit.  
The deposit resulted in a penalty reduction, and the bank reduced the penalty reduction.  
Is there an event that led to the deposit?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both segments describe the same sequence of events: John violated the overdraft policy, a fee was charged, that charge led to a deposit, and the deposit caused the bank to reduce a penalty. The main differences are stylistic: the second text omits “bank” in the first sentence, rephrases “charging of a fee” as “fee charging,” and adds a redundant phrase “the bank reduced the penalty reduction,” which slightly alters clarity but not the core meaning. The question at the end is also phrased slightly differently but asks essentially the same thing.
% -------------------------------------------------------------------------
% The final logic is :
tff(type_event,type,
    event: $tType ).

tff(type_entity,type,
    entity: $tType ).

tff(type_john,type,
    john: entity ).

tff(type_overdraft_policy,type,
    overdraft_policy: entity ).

tff(type_fee_charging,type,
    fee_charging: event ).

tff(type_deposit,type,
    deposit: event ).

tff(type_penalty_reduction,type,
    penalty_reduction: event ).

tff(type_bank,type,
    bank: entity ).

tff(type_in_noncompliance_with,type,
    in_noncompliance_with: ( entity * entity ) > $o ).

tff(type_led_to,type,
    led_to: ( event * event ) > $o ).

tff(type_resulted_in,type,
    resulted_in: ( event * event ) > $o ).

tff(type_reduced,type,
    reduced: ( entity * event ) > $o ).

tff(ax1,axiom,
    in_noncompliance_with(john,overdraft_policy) ).

tff(ax2,axiom,
    led_to(fee_charging,deposit) ).

tff(ax3,axiom,
    ( resulted_in(deposit,penalty_reduction)
    & reduced(bank,penalty_reduction) ) ).

tff(conj1,conjecture,
    ? [E: event] : led_to(E,deposit) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, event: $tType).
tff(type_def_6, type, entity: $tType).
tff(func_def_0, type, john: entity).
tff(func_def_1, type, overdraft_policy: entity).
tff(func_def_2, type, fee_charging: event).
tff(func_def_3, type, deposit: event).
tff(func_def_4, type, penalty_reduction: event).
tff(func_def_5, type, bank: entity).
tff(pred_def_1, type, in_noncompliance_with: (entity * entity) > $o).
tff(pred_def_2, type, led_to: (event * event) > $o).
tff(pred_def_3, type, resulted_in: (event * event) > $o).
tff(pred_def_4, type, reduced: (entity * event) > $o).
tff(f12,plain,(
  $false),
  inference(resolution,[],[f10,f11])).
tff(f11,plain,(
  ( ! [X0 : event] : (~led_to(X0,deposit)) )),
  inference(cnf_transformation,[],[f9])).
tff(f9,plain,(
  ! [X0 : event] : ~led_to(X0,deposit)),
  inference(ennf_transformation,[],[f5])).
tff(f5,negated_conjecture,(
  ~? [X0 : event] : led_to(X0,deposit)),
  inference(negated_conjecture,[status(cth)],[f4])).
tff(f4,conjecture,(
  ? [X0 : event] : led_to(X0,deposit)),
  file('/tmp/dn5n_LxZNx/SOT_fzgy7X',unknown)).
tff(f10,plain,(
  led_to(fee_charging,deposit)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  led_to(fee_charging,deposit)),
  file('/tmp/dn5n_LxZNx/SOT_fzgy7X',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-AP/02-Bank.THM.txt,False,0.93,3,True,1.00,2,THM,Unknown,0,0,0,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 3:1 with similarity 0.97
% The original and final NL have similarity: 0.98
% NL2L model:       Google---gemini-2.5-flash
% L2NL model:       Google---gemini-2.5-flash
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
John was in noncompliance with the bank overdraft policy.
The charging of a fee led to a deposit.
As a result of the deposit, the bank reduced the penalty.
What led to a deposit?

% -------------------------------------------------------------------------
% The final NL is :
John is in noncompliance with the bank overdraft policy.
Charging of a fee led to a deposit.
As a result of a deposit, the bank reduced the penalty.
What led to the deposit?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
The two segments differ only in minor grammatical details (“was” vs. “is,” “the deposit” vs. “a deposit”) and article usage. These do not change the meaning: both describe the same sequence of events and ask the same question.
% -------------------------------------------------------------------------
% The final logic is :
tff(type_john,type,
    john: $i ).

tff(type_noncompliance,type,
    noncompliance: $i ).

tff(type_bank_overdraft_policy,type,
    bank_overdraft_policy: $i ).

tff(type_in_noncompliance_with,type,
    in_noncompliance_with: ( $i * $i * $i ) > $o ).

tff(type_charging_of_a_fee,type,
    charging_of_a_fee: $i ).

tff(type_deposit,type,
    deposit: $i ).

tff(type_led_to,type,
    led_to: ( $i * $i ) > $o ).

tff(type_bank,type,
    bank: $i ).

tff(type_penalty,type,
    penalty: $i ).

tff(type_reduced,type,
    reduced: ( $i * $i ) > $o ).

tff(type_as_a_result_of,type,
    as_a_result_of: ( $i * $o ) > $o ).

tff(john_in_noncompliance_with_policy,axiom,
    in_noncompliance_with(john,noncompliance,bank_overdraft_policy) ).

tff(charging_fee_led_to_deposit,axiom,
    led_to(charging_of_a_fee,deposit) ).

tff(deposit_led_to_penalty_reduction,axiom,
    as_a_result_of(deposit,reduced(bank,penalty)) ).

tff(what_led_to_deposit,conjecture,
    ? [X] : led_to(X,deposit) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(pred_def_4, type, as_a_result_of: ($i * $o) > $o).
fof(f11,plain,(
  $false),
  inference(subsumption_resolution,[],[f9,f10])).
fof(f10,plain,(
  ( ! [X0] : (~led_to(X0,deposit)) )),
  inference(cnf_transformation,[],[f8])).
fof(f8,plain,(
  ! [X0] : ~led_to(X0,deposit)),
  inference(ennf_transformation,[],[f5])).
fof(f5,negated_conjecture,(
  ~? [X0] : led_to(X0,deposit)),
  inference(negated_conjecture,[status(cth)],[f4])).
fof(f4,conjecture,(
  ? [X0] : led_to(X0,deposit)),
  file('/tmp/DcZYr5Kf5p/SOT_LVoMiD',unknown)).
fof(f9,plain,(
  led_to(charging_of_a_fee,deposit)),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  led_to(charging_of_a_fee,deposit)),
  file('/tmp/DcZYr5Kf5p/SOT_LVoMiD',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-AP/02-Bank.THM.txt,True,0.98,3,True,0.97,1,THM,Unknown,1,0,0,Google---gemini-2.5-flash,Google---gemini-2.5-flash,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 1:1 with similarity 0.95
% The original and final NL have similarity: 0.95
% NL2L model:       Google---gemini-2.5-flash
% L2NL model:       Google---gemini-2.5-flash
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
John was in noncompliance with the bank overdraft policy.
The charging of a fee led to a deposit.
As a result of the deposit, the bank reduced the penalty.
What led to a deposit?

% -------------------------------------------------------------------------
% The final NL is :
'john' is an individual.
'noncompliance' is an individual.
'bank_overdraft_policy' is an individual.
'charging_of_a_fee' is an individual.
'deposit' is an individual.
'bank_reducing_penalty' is an individual.
'was_in' is a binary predicate.
'with' is a binary predicate.
'led_to' is a binary predicate.
'as_a_result_of' is a binary predicate.
John was in noncompliance with the bank overdraft policy.
Charging of a fee led to a deposit.
The bank reducing penalty was as a result of the deposit.
What led to a deposit?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both segments describe the same scenario and ask the same question: a fee being charged led to a deposit. The second segment, however, adds explicit logical or ontological labeling (declaring individuals and predicates) before restating the same content. This makes it more formal and structured but does not change the underlying meaning.
% -------------------------------------------------------------------------
% The final logic is :
tff(john_type,type,
    john: $i ).

tff(noncompliance_type,type,
    noncompliance: $i ).

tff(bank_overdraft_policy_type,type,
    bank_overdraft_policy: $i ).

tff(charging_of_a_fee_type,type,
    charging_of_a_fee: $i ).

tff(deposit_type,type,
    deposit: $i ).

tff(bank_reducing_penalty_type,type,
    bank_reducing_penalty: $i ).

tff(was_in_type,type,
    was_in: ( $i * $i ) > $o ).

tff(with_type,type,
    with: ( $i * $i ) > $o ).

tff(led_to_type,type,
    led_to: ( $i * $i ) > $o ).

tff(as_a_result_of_type,type,
    as_a_result_of: ( $i * $i ) > $o ).

tff(john_was_in_noncompliance_with_bank_overdraft_policy,axiom,
    ( was_in(john,noncompliance)
    & with(noncompliance,bank_overdraft_policy) ) ).

tff(charging_of_a_fee_led_to_a_deposit,axiom,
    led_to(charging_of_a_fee,deposit) ).

tff(as_a_result_of_the_deposit_the_bank_reduced_the_penalty,axiom,
    as_a_result_of(deposit,bank_reducing_penalty) ).

tff(what_led_to_a_deposit,conjecture,
    ? [X] : led_to(X,deposit) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
fof(f12,plain,(
  $false),
  inference(subsumption_resolution,[],[f10,f11])).
fof(f11,plain,(
  ( ! [X0] : (~led_to(X0,deposit)) )),
  inference(cnf_transformation,[],[f9])).
fof(f9,plain,(
  ! [X0] : ~led_to(X0,deposit)),
  inference(ennf_transformation,[],[f5])).
fof(f5,negated_conjecture,(
  ~? [X0] : led_to(X0,deposit)),
  inference(negated_conjecture,[status(cth)],[f4])).
fof(f4,conjecture,(
  ? [X0] : led_to(X0,deposit)),
  file('/tmp/PllNm7Kjhy/SOT_d3AsIW',unknown)).
fof(f10,plain,(
  led_to(charging_of_a_fee,deposit)),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  led_to(charging_of_a_fee,deposit)),
  file('/tmp/PllNm7Kjhy/SOT_d3AsIW',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-AP/02-Bank.THM.txt,True,0.95,1,True,0.95,1,THM,Unknown,2,0,0,Google---gemini-2.5-flash,Google---gemini-2.5-flash,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 1:1 with similarity 0.95
% The original and final NL have similarity: 0.95
% NL2L model:       Google---gemini-2.5-flash
% L2NL model:       Google---gemini-2.5-flash
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
John was in noncompliance with the bank overdraft policy.
The charging of a fee led to a deposit.
As a result of the deposit, the bank reduced the penalty.
What led to a deposit?

% -------------------------------------------------------------------------
% The final NL is :
Person is a type.
Policy is a type.
Event is a type.
John is a person.
Bank Overdraft Policy is a policy.
Charging of Fee is an event.
Deposit is an event.
Reduction of Penalty is an event.
'In non-compliance' is a relation between a person and a policy.
'Causes' is a relation between two events.
John is in non-compliance with the Bank Overdraft Policy.
Charging of Fee causes a Deposit.
A Deposit causes a Reduction of Penalty.
Is there an event X such that X causes a Deposit?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both segments describe the same situation: John violates a bank policy, a fee is charged, that fee leads to a deposit, and the deposit leads to a penalty reduction. The first text expresses this narratively and asks what led to a deposit; the second formalizes the same facts in a structured, ontology-like form and asks whether there exists an event that causes a deposit. The only difference is in style and explicitness—the second adds type declarations and relation definitions, but the underlying meaning and causal chain are nearly identical.
% -------------------------------------------------------------------------
% The final logic is :
tff(type_person,type,
    person: $tType ).

tff(type_policy,type,
    policy: $tType ).

tff(type_event,type,
    event: $tType ).

tff(decl_john,type,
    john: person ).

tff(decl_bankOverdraftPolicy,type,
    bankOverdraftPolicy: policy ).

tff(decl_chargingOfFee,type,
    chargingOfFee: event ).

tff(decl_deposit,type,
    deposit: event ).

tff(decl_reductionOfPenalty,type,
    reductionOfPenalty: event ).

tff(decl_inNoncompliance,type,
    inNoncompliance: ( person * policy ) > $o ).

tff(decl_causes,type,
    causes: ( event * event ) > $o ).

tff(sentence1,axiom,
    inNoncompliance(john,bankOverdraftPolicy) ).

tff(sentence2,axiom,
    causes(chargingOfFee,deposit) ).

tff(sentence3,axiom,
    causes(deposit,reductionOfPenalty) ).

tff(question1,conjecture,
    ? [X: event] : causes(X,deposit) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, person: $tType).
tff(type_def_6, type, policy: $tType).
tff(type_def_7, type, event: $tType).
tff(func_def_0, type, john: person).
tff(func_def_1, type, bankOverdraftPolicy: policy).
tff(func_def_2, type, chargingOfFee: event).
tff(func_def_3, type, deposit: event).
tff(func_def_4, type, reductionOfPenalty: event).
tff(pred_def_1, type, inNoncompliance: (person * policy) > $o).
tff(pred_def_2, type, causes: (event * event) > $o).
tff(f11,plain,(
  $false),
  inference(subsumption_resolution,[],[f8,f10])).
tff(f10,plain,(
  ( ! [X0 : event] : (~causes(X0,deposit)) )),
  inference(cnf_transformation,[],[f7])).
tff(f7,plain,(
  ! [X0 : event] : ~causes(X0,deposit)),
  inference(ennf_transformation,[],[f5])).
tff(f5,negated_conjecture,(
  ~? [X0 : event] : causes(X0,deposit)),
  inference(negated_conjecture,[status(cth)],[f4])).
tff(f4,conjecture,(
  ? [X0 : event] : causes(X0,deposit)),
  file('/tmp/gu9GPDd_OH/SOT_prxzUI',unknown)).
tff(f8,plain,(
  causes(chargingOfFee,deposit)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  causes(chargingOfFee,deposit)),
  file('/tmp/gu9GPDd_OH/SOT_prxzUI',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-AP/02-Bank.THM.txt,True,0.95,1,True,0.95,1,THM,Unknown,1,1,0,Google---gemini-2.5-flash,Google---gemini-2.5-flash,OpenAI---gpt-5-chat-latest
