% -------------------------------------------------------------------------
% Converged at ZigZag number 1:1 with similarity 0.95
% The original and final NL have similarity: 0.95
% NL2L model:       OpenAI---gpt-5-chat-latest
% L2NL model:       OpenAI---gpt-5-chat-latest
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
If a country is a member of NATO, it will protect any member that is attacked.
Sweden is a member of NATO.
Germany is a member of NATO
Russia attacked Germany.
Will Sweden protect Germany?

% -------------------------------------------------------------------------
% The final NL is :
There is a type of thing called a country.  

“Member” is a relation between two countries.  

“Attacked” is a relation between two countries.  

“Protects” is a relation between two countries.  

NATO is a country.  

Sweden is a country.  

Germany is a country.  

Russia is a country.  

For every country C, if C is a member of NATO, then for all countries A and M, if M is a member of NATO and A has attacked M, then C protects M.  

Sweden is a member of NATO.  

Germany is a member of NATO.  

Russia has attacked Germany.  

Does Sweden protect Germany?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both texts describe the same logical situation: NATO membership implies mutual protection among members when one is attacked, and they ask whether Sweden will protect Germany after Russia attacks Germany. The second text is more formal and explicit about the logical structure (defining relations and quantifiers), while the first is a natural-language version. The only minor difference is that the second explicitly defines “NATO” as a country and formalizes the relations, but the underlying meaning and question are effectively identical.
% -------------------------------------------------------------------------
% The final logic is :
tff(type_country,type,
    country: $tType ).

tff(type_member,type,
    member: ( country * country ) > $o ).

tff(type_attacked,type,
    attacked: ( country * country ) > $o ).

tff(type_protects,type,
    protects: ( country * country ) > $o ).

tff(type_nato,type,
    nato: country ).

tff(type_sweden,type,
    sweden: country ).

tff(type_germany,type,
    germany: country ).

tff(type_russia,type,
    russia: country ).

tff(axiom_nato_protection,axiom,
    ! [C: country] :
      ( member(C,nato)
     => ! [A: country,M: country] :
          ( ( member(M,nato)
            & attacked(A,M) )
         => protects(C,M) ) ) ).

tff(axiom_sweden_member,axiom,
    member(sweden,nato) ).

tff(axiom_germany_member,axiom,
    member(germany,nato) ).

tff(axiom_russia_attacked_germany,axiom,
    attacked(russia,germany) ).

tff(conjecture_sweden_protects_germany,conjecture,
    protects(sweden,germany) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, country: $tType).
tff(func_def_0, type, nato: country).
tff(func_def_1, type, sweden: country).
tff(func_def_2, type, germany: country).
tff(func_def_3, type, russia: country).
tff(pred_def_1, type, member: (country * country) > $o).
tff(pred_def_2, type, attacked: (country * country) > $o).
tff(pred_def_3, type, protects: (country * country) > $o).
tff(pred_def_4, type, sP0: country > $o).
tff(pred_def_5, type, sP1: country > $o).
tff(pred_def_6, type, sP2: country > $o).
tff(f26,plain,(
  $false),
  inference(finite_model_not_found_:_exhaustively_excluded_all_possible_domain_size_assignments,[],[f10,f11,f12,f13,f14])).
tff(f14,plain,(
  ~protects(sweden,germany)),
  inference(cnf_transformation,[],[f7])).
tff(f7,plain,(
  ~protects(sweden,germany)),
  inference(flattening,[],[f6])).
tff(f6,negated_conjecture,(
  ~protects(sweden,germany)),
  inference(negated_conjecture,[status(cth)],[f5])).
tff(f5,conjecture,(
  protects(sweden,germany)),
  file('/tmp/UzZgHPkFWO/SOT_awcSxt',unknown)).
tff(f13,plain,(
  attacked(russia,germany)),
  inference(cnf_transformation,[],[f4])).
tff(f4,axiom,(
  attacked(russia,germany)),
  file('/tmp/UzZgHPkFWO/SOT_awcSxt',unknown)).
tff(f12,plain,(
  member(germany,nato)),
  inference(cnf_transformation,[],[f3])).
tff(f3,axiom,(
  member(germany,nato)),
  file('/tmp/UzZgHPkFWO/SOT_awcSxt',unknown)).
tff(f11,plain,(
  member(sweden,nato)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  member(sweden,nato)),
  file('/tmp/UzZgHPkFWO/SOT_awcSxt',unknown)).
tff(f10,plain,(
  ( ! [X2 : country,X0 : country,X1 : country] : (protects(X0,X2) | ~attacked(X1,X2) | ~member(X2,nato) | ~member(X0,nato)) )),
  inference(cnf_transformation,[],[f9])).
tff(f9,plain,(
  ! [X0 : country] : (! [X1 : country,X2 : country] : (protects(X0,X2) | ~attacked(X1,X2) | ~member(X2,nato)) | ~member(X0,nato))),
  inference(flattening,[],[f8])).
tff(f8,plain,(
  ! [X0 : country] : (! [X1 : country,X2 : country] : (protects(X0,X2) | (~attacked(X1,X2) | ~member(X2,nato))) | ~member(X0,nato))),
  inference(ennf_transformation,[],[f1])).
tff(f1,axiom,(
  ! [X0 : country] : (member(X0,nato) => ! [X1 : country,X2 : country] : ((attacked(X1,X2) & member(X2,nato)) => protects(X0,X2)))),
  file('/tmp/UzZgHPkFWO/SOT_awcSxt',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-RT/NonTrivial_03.THM.txt,True,0.95,1,True,0.95,1,THM,Unknown,1,0,0,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 1:1 with similarity 0.95
% The original and final NL have similarity: 0.95
% NL2L model:       OpenAI---gpt-5-chat-latest
% L2NL model:       OpenAI---gpt-5-chat-latest
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
If a country is a member of NATO, it will protect any member that is attacked.
Sweden is a member of NATO.
Germany is a member of NATO
Russia attacked Germany.
Will Sweden protect Germany?

% -------------------------------------------------------------------------
% The final NL is :
For every country, if that country is a member of NATO, then for all countries M and A, if M is a member of NATO and A attacks M, then that country protects M.  

Sweden is a member of NATO.  

Germany is a member of NATO.  

Russia attacked Germany.  

Does Sweden protect Germany?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both segments express the same logical structure: NATO membership implies mutual defense among members, and the scenario tests whether Sweden will protect Germany after Russia attacks Germany. The second segment is more formal and generalized, using quantifiers (“for every country,” “for all countries M and A”) to express the rule, while the first uses a more natural-language conditional. The meaning, however, is nearly identical.
% -------------------------------------------------------------------------
% The final logic is :
tff(type_country,type,
    country: $tType ).

tff(type_organization,type,
    organization: $tType ).

tff(type_member,type,
    member: ( country * organization ) > $o ).

tff(type_attacked,type,
    attacked: ( country * country ) > $o ).

tff(type_protects,type,
    protects: ( country * country ) > $o ).

tff(type_nato,type,
    nato: organization ).

tff(type_sweden,type,
    sweden: country ).

tff(type_germany,type,
    germany: country ).

tff(type_russia,type,
    russia: country ).

tff(ax1,axiom,
    ! [C: country] :
      ( member(C,nato)
     => ! [M: country,A: country] :
          ( ( member(M,nato)
            & attacked(A,M) )
         => protects(C,M) ) ) ).

tff(ax2,axiom,
    member(sweden,nato) ).

tff(ax3,axiom,
    member(germany,nato) ).

tff(ax4,axiom,
    attacked(russia,germany) ).

tff(conj,conjecture,
    protects(sweden,germany) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, country: $tType).
tff(type_def_6, type, organization: $tType).
tff(func_def_0, type, nato: organization).
tff(func_def_1, type, sweden: country).
tff(func_def_2, type, germany: country).
tff(func_def_3, type, russia: country).
tff(pred_def_1, type, member: (country * organization) > $o).
tff(pred_def_2, type, attacked: (country * country) > $o).
tff(pred_def_3, type, protects: (country * country) > $o).
tff(f39,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f26,f36,f38])).
tff(f38,plain,(
  ~spl0_1),
  inference(avatar_contradiction_clause,[],[f37])).
tff(f37,plain,(
  $false | ~spl0_1),
  inference(resolution,[],[f22,f14])).
tff(f14,plain,(
  ~protects(sweden,germany)),
  inference(cnf_transformation,[],[f7])).
tff(f7,plain,(
  ~protects(sweden,germany)),
  inference(flattening,[],[f6])).
tff(f6,negated_conjecture,(
  ~protects(sweden,germany)),
  inference(negated_conjecture,[status(cth)],[f5])).
tff(f5,conjecture,(
  protects(sweden,germany)),
  file('/tmp/lTIRAUHYGP/SOT_kg1SWY',unknown)).
tff(f22,plain,(
  protects(sweden,germany) | ~spl0_1),
  inference(avatar_component_clause,[],[f20])).
tff(f20,definition,(
  spl0_1 <=> protects(sweden,germany)),
  introduced(definition,[new_symbols(naming,[spl0_1])],[avatar_definition])).
tff(f36,plain,(
  ~spl0_2),
  inference(avatar_contradiction_clause,[],[f35])).
tff(f35,plain,(
  $false | ~spl0_2),
  inference(resolution,[],[f25,f13])).
tff(f13,plain,(
  attacked(russia,germany)),
  inference(cnf_transformation,[],[f4])).
tff(f4,axiom,(
  attacked(russia,germany)),
  file('/tmp/lTIRAUHYGP/SOT_kg1SWY',unknown)).
tff(f25,plain,(
  ( ! [X0 : country] : (~attacked(X0,germany)) ) | ~spl0_2),
  inference(avatar_component_clause,[],[f24])).
tff(f24,definition,(
  spl0_2 <=> ! [X0 : country] : ~attacked(X0,germany)),
  introduced(definition,[new_symbols(naming,[spl0_2])],[avatar_definition])).
tff(f26,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f18,f24,f20])).
tff(f18,plain,(
  ( ! [X0 : country] : (~attacked(X0,germany) | protects(sweden,germany)) )),
  inference(resolution,[],[f15,f12])).
tff(f12,plain,(
  member(germany,nato)),
  inference(cnf_transformation,[],[f3])).
tff(f3,axiom,(
  member(germany,nato)),
  file('/tmp/lTIRAUHYGP/SOT_kg1SWY',unknown)).
tff(f15,plain,(
  ( ! [X0 : country,X1 : country] : (~member(X1,nato) | ~attacked(X0,X1) | protects(sweden,X1)) )),
  inference(resolution,[],[f10,f11])).
tff(f11,plain,(
  member(sweden,nato)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  member(sweden,nato)),
  file('/tmp/lTIRAUHYGP/SOT_kg1SWY',unknown)).
tff(f10,plain,(
  ( ! [X2 : country,X0 : country,X1 : country] : (~member(X0,nato) | ~attacked(X2,X1) | ~member(X1,nato) | protects(X0,X1)) )),
  inference(cnf_transformation,[],[f9])).
tff(f9,plain,(
  ! [X0 : country] : (! [X1 : country,X2 : country] : (protects(X0,X1) | ~attacked(X2,X1) | ~member(X1,nato)) | ~member(X0,nato))),
  inference(flattening,[],[f8])).
tff(f8,plain,(
  ! [X0 : country] : (! [X1 : country,X2 : country] : (protects(X0,X1) | (~attacked(X2,X1) | ~member(X1,nato))) | ~member(X0,nato))),
  inference(ennf_transformation,[],[f1])).
tff(f1,axiom,(
  ! [X0 : country] : (member(X0,nato) => ! [X1 : country,X2 : country] : ((attacked(X2,X1) & member(X1,nato)) => protects(X0,X1)))),
  file('/tmp/lTIRAUHYGP/SOT_kg1SWY',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-RT/NonTrivial_03.THM.txt,True,0.95,1,True,0.95,1,THM,Unknown,1,0,0,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 1:1 with similarity 0.95
% The original and final NL have similarity: 0.95
% NL2L model:       OpenAI---gpt-5-chat-latest
% L2NL model:       OpenAI---gpt-5-chat-latest
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
If a country is a member of NATO, it will protect any member that is attacked.
Sweden is a member of NATO.
Germany is a member of NATO
Russia attacked Germany.
Will Sweden protect Germany?

% -------------------------------------------------------------------------
% The final NL is :
For every country, if that country is a member of NATO, then for every country M, if M is a member of NATO and Russia attacks M, that country protects M.  

Sweden is a member of NATO.  

Germany is a member of NATO.  

Russia attacked Germany.  

Does Sweden protect Germany?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both segments express the same logical structure: NATO membership implies mutual protection among members if one is attacked. The second text is a more formal, generalized version using quantifiers (“for every country”), while the first uses a more natural, narrative phrasing. The meaning is nearly identical, differing only in style and explicit logical formulation.
% -------------------------------------------------------------------------
% The final logic is :
tff(type_country,type,
    country: $tType ).

tff(type_member,type,
    member: ( country * country ) > $o ).

tff(type_attacked,type,
    attacked: ( country * country ) > $o ).

tff(type_protects,type,
    protects: ( country * country ) > $o ).

tff(type_nato,type,
    nato: country ).

tff(type_sweden,type,
    sweden: country ).

tff(type_germany,type,
    germany: country ).

tff(type_russia,type,
    russia: country ).

tff(axiom1,axiom,
    ! [C: country] :
      ( member(C,nato)
     => ! [M: country] :
          ( ( member(M,nato)
            & attacked(russia,M) )
         => protects(C,M) ) ) ).

tff(axiom2,axiom,
    member(sweden,nato) ).

tff(axiom3,axiom,
    member(germany,nato) ).

tff(axiom4,axiom,
    attacked(russia,germany) ).

tff(conjecture1,conjecture,
    protects(sweden,germany) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, country: $tType).
tff(func_def_0, type, nato: country).
tff(func_def_1, type, sweden: country).
tff(func_def_2, type, germany: country).
tff(func_def_3, type, russia: country).
tff(pred_def_1, type, member: (country * country) > $o).
tff(pred_def_2, type, attacked: (country * country) > $o).
tff(pred_def_3, type, protects: (country * country) > $o).
tff(f19,plain,(
  $false),
  inference(subsumption_resolution,[],[f17,f14])).
tff(f14,plain,(
  ~protects(sweden,germany)),
  inference(cnf_transformation,[],[f7])).
tff(f7,plain,(
  ~protects(sweden,germany)),
  inference(flattening,[],[f6])).
tff(f6,negated_conjecture,(
  ~protects(sweden,germany)),
  inference(negated_conjecture,[status(cth)],[f5])).
tff(f5,conjecture,(
  protects(sweden,germany)),
  file('/tmp/scKtVPturJ/SOT_3INJ71',unknown)).
tff(f17,plain,(
  protects(sweden,germany)),
  inference(resolution,[],[f16,f11])).
tff(f11,plain,(
  member(sweden,nato)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  member(sweden,nato)),
  file('/tmp/scKtVPturJ/SOT_3INJ71',unknown)).
tff(f16,plain,(
  ( ! [X0 : country] : (~member(X0,nato) | protects(X0,germany)) )),
  inference(subsumption_resolution,[],[f15,f12])).
tff(f12,plain,(
  member(germany,nato)),
  inference(cnf_transformation,[],[f3])).
tff(f3,axiom,(
  member(germany,nato)),
  file('/tmp/scKtVPturJ/SOT_3INJ71',unknown)).
tff(f15,plain,(
  ( ! [X0 : country] : (protects(X0,germany) | ~member(germany,nato) | ~member(X0,nato)) )),
  inference(resolution,[],[f10,f13])).
tff(f13,plain,(
  attacked(russia,germany)),
  inference(cnf_transformation,[],[f4])).
tff(f4,axiom,(
  attacked(russia,germany)),
  file('/tmp/scKtVPturJ/SOT_3INJ71',unknown)).
tff(f10,plain,(
  ( ! [X0 : country,X1 : country] : (~attacked(russia,X1) | protects(X0,X1) | ~member(X1,nato) | ~member(X0,nato)) )),
  inference(cnf_transformation,[],[f9])).
tff(f9,plain,(
  ! [X0 : country] : (! [X1 : country] : (protects(X0,X1) | ~attacked(russia,X1) | ~member(X1,nato)) | ~member(X0,nato))),
  inference(flattening,[],[f8])).
tff(f8,plain,(
  ! [X0 : country] : (! [X1 : country] : (protects(X0,X1) | (~attacked(russia,X1) | ~member(X1,nato))) | ~member(X0,nato))),
  inference(ennf_transformation,[],[f1])).
tff(f1,axiom,(
  ! [X0 : country] : (member(X0,nato) => ! [X1 : country] : ((attacked(russia,X1) & member(X1,nato)) => protects(X0,X1)))),
  file('/tmp/scKtVPturJ/SOT_3INJ71',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-RT/NonTrivial_03.THM.txt,True,0.95,1,True,0.95,1,THM,Unknown,2,0,0,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 1:1 with similarity 0.95
% The original and final NL have similarity: 0.95
% NL2L model:       Google---gemini-2.5-flash
% L2NL model:       Google---gemini-2.5-flash
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
If a country is a member of NATO, it will protect any member that is attacked.
Sweden is a member of NATO.
Germany is a member of NATO
Russia attacked Germany.
Will Sweden protect Germany?

% -------------------------------------------------------------------------
% The final NL is :
"country" is a type.
"member_of_NATO" is a predicate that applies to a country.
"protects" is a predicate that applies to two countries.
"attacked" is a predicate that applies to a country.
Sweden is a country.
Germany is a country.
Russia is a country.
For every country C and every country M, if C is a member of NATO, then if M is a member of NATO and M is attacked, C protects M.
Sweden is a member of NATO.
Germany is a member of NATO.
Germany is attacked.
Does Sweden protect Germany?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both segments express the same logical content: NATO members protect other members that are attacked, and the question is whether Sweden will protect Germany after Russia attacks Germany.  
The first segment uses natural language, while the second formalizes the same information in logical or predicate form, explicitly defining types and predicates. The only minor difference is that the second text is more explicit and formal, but the meaning conveyed is essentially identical.
% -------------------------------------------------------------------------
% The final logic is :
tff(country_type,type,
    country: $tType ).

tff(member_of_NATO_type,type,
    member_of_NATO: country > $o ).

tff(protects_type,type,
    protects: ( country * country ) > $o ).

tff(attacked_type,type,
    attacked: country > $o ).

tff(sweden_declaration,type,
    sweden: country ).

tff(germany_declaration,type,
    germany: country ).

tff(russia_declaration,type,
    russia: country ).

tff(axiom_if_country_is_member_of_NATO_it_protects_any_member_that_is_attacked,axiom,
    ! [C: country,M: country] :
      ( member_of_NATO(C)
     => ( ( member_of_NATO(M)
          & attacked(M) )
       => protects(C,M) ) ) ).

tff(axiom_sweden_is_a_member_of_NATO,axiom,
    member_of_NATO(sweden) ).

tff(axiom_germany_is_a_member_of_NATO,axiom,
    member_of_NATO(germany) ).

tff(axiom_russia_attacked_germany,axiom,
    attacked(germany) ).

tff(conjecture_will_sweden_protect_germany,conjecture,
    protects(sweden,germany) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, country: $tType).
tff(func_def_0, type, sweden: country).
tff(func_def_1, type, germany: country).
tff(func_def_2, type, russia: country).
tff(pred_def_1, type, member_of_NATO: country > $o).
tff(pred_def_2, type, protects: (country * country) > $o).
tff(pred_def_3, type, attacked: country > $o).
tff(f19,plain,(
  $false),
  inference(finite_model_not_found_:_exhaustively_excluded_all_possible_domain_size_assignments,[],[f10,f11,f12,f13,f14])).
tff(f14,plain,(
  ~protects(sweden,germany)),
  inference(cnf_transformation,[],[f7])).
tff(f7,plain,(
  ~protects(sweden,germany)),
  inference(flattening,[],[f6])).
tff(f6,negated_conjecture,(
  ~protects(sweden,germany)),
  inference(negated_conjecture,[status(cth)],[f5])).
tff(f5,conjecture,(
  protects(sweden,germany)),
  file('/tmp/q3BIAXEsX2/SOT_rgmBHz',unknown)).
tff(f13,plain,(
  attacked(germany)),
  inference(cnf_transformation,[],[f4])).
tff(f4,axiom,(
  attacked(germany)),
  file('/tmp/q3BIAXEsX2/SOT_rgmBHz',unknown)).
tff(f12,plain,(
  member_of_NATO(germany)),
  inference(cnf_transformation,[],[f3])).
tff(f3,axiom,(
  member_of_NATO(germany)),
  file('/tmp/q3BIAXEsX2/SOT_rgmBHz',unknown)).
tff(f11,plain,(
  member_of_NATO(sweden)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  member_of_NATO(sweden)),
  file('/tmp/q3BIAXEsX2/SOT_rgmBHz',unknown)).
tff(f10,plain,(
  ( ! [X0 : country,X1 : country] : (protects(X0,X1) | ~attacked(X1) | ~member_of_NATO(X1) | ~member_of_NATO(X0)) )),
  inference(cnf_transformation,[],[f9])).
tff(f9,plain,(
  ! [X0 : country,X1 : country] : (protects(X0,X1) | ~attacked(X1) | ~member_of_NATO(X1) | ~member_of_NATO(X0))),
  inference(flattening,[],[f8])).
tff(f8,plain,(
  ! [X0 : country,X1 : country] : ((protects(X0,X1) | (~attacked(X1) | ~member_of_NATO(X1))) | ~member_of_NATO(X0))),
  inference(ennf_transformation,[],[f1])).
tff(f1,axiom,(
  ! [X0 : country,X1 : country] : (member_of_NATO(X0) => ((attacked(X1) & member_of_NATO(X1)) => protects(X0,X1)))),
  file('/tmp/q3BIAXEsX2/SOT_rgmBHz',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-RT/NonTrivial_03.THM.txt,True,0.95,1,True,0.95,1,THM,Unknown,1,1,0,Google---gemini-2.5-flash,Google---gemini-2.5-flash,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 1:1 with similarity 0.95
% The original and final NL have similarity: 0.95
% NL2L model:       Google---gemini-2.5-flash
% L2NL model:       Google---gemini-2.5-flash
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
If a country is a member of NATO, it will protect any member that is attacked.
Sweden is a member of NATO.
Germany is a member of NATO
Russia attacked Germany.
Will Sweden protect Germany?

% -------------------------------------------------------------------------
% The final NL is :
Country is a type.
Sweden is a country.
Germany is a country.
Russia is a country.
Member of NATO is a property of countries.
Attacked is a property of countries.
Protects is a relation between two countries.
For any country C and any country M, if C is a member of NATO, and M is a member of NATO, and M is attacked, then C protects M.
Sweden is a member of NATO.
Germany is a member of NATO.
Germany is attacked.
Does Sweden protect Germany?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both segments express the same logical content: that NATO membership entails mutual protection among members, and they apply this to Sweden and Germany when Germany is attacked.  
The first text uses natural language and a narrative form, while the second formalizes the same information in logical or ontological terms (defining types, properties, and relations). The only minor difference is that the second explicitly defines the logical structure and predicates, while the first assumes them implicitly.
% -------------------------------------------------------------------------
% The final logic is :
tff(country_type,type,
    country: $tType ).

tff(sweden_type,type,
    sweden: country ).

tff(germany_type,type,
    germany: country ).

tff(russia_type,type,
    russia: country ).

tff(member_of_nato_type,type,
    member_of_nato: country > $o ).

tff(attacked_type,type,
    attacked: country > $o ).

tff(protects_type,type,
    protects: ( country * country ) > $o ).

tff(s1,axiom,
    ! [C: country,M: country] :
      ( ( member_of_nato(C)
        & member_of_nato(M)
        & attacked(M) )
     => protects(C,M) ) ).

tff(s2,axiom,
    member_of_nato(sweden) ).

tff(s3,axiom,
    member_of_nato(germany) ).

tff(s4,axiom,
    attacked(germany) ).

tff(s5,conjecture,
    protects(sweden,germany) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, country: $tType).
tff(func_def_0, type, sweden: country).
tff(func_def_1, type, germany: country).
tff(func_def_2, type, russia: country).
tff(pred_def_1, type, member_of_nato: country > $o).
tff(pred_def_2, type, attacked: country > $o).
tff(pred_def_3, type, protects: (country * country) > $o).
tff(f19,plain,(
  $false),
  inference(subsumption_resolution,[],[f17,f14])).
tff(f14,plain,(
  ~protects(sweden,germany)),
  inference(cnf_transformation,[],[f7])).
tff(f7,plain,(
  ~protects(sweden,germany)),
  inference(flattening,[],[f6])).
tff(f6,negated_conjecture,(
  ~protects(sweden,germany)),
  inference(negated_conjecture,[status(cth)],[f5])).
tff(f5,conjecture,(
  protects(sweden,germany)),
  file('/tmp/_tuzWjGqvL/SOT_E5Zjwa',unknown)).
tff(f17,plain,(
  protects(sweden,germany)),
  inference(resolution,[],[f16,f11])).
tff(f11,plain,(
  member_of_nato(sweden)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  member_of_nato(sweden)),
  file('/tmp/_tuzWjGqvL/SOT_E5Zjwa',unknown)).
tff(f16,plain,(
  ( ! [X0 : country] : (~member_of_nato(X0) | protects(X0,germany)) )),
  inference(subsumption_resolution,[],[f15,f12])).
tff(f12,plain,(
  member_of_nato(germany)),
  inference(cnf_transformation,[],[f3])).
tff(f3,axiom,(
  member_of_nato(germany)),
  file('/tmp/_tuzWjGqvL/SOT_E5Zjwa',unknown)).
tff(f15,plain,(
  ( ! [X0 : country] : (protects(X0,germany) | ~member_of_nato(germany) | ~member_of_nato(X0)) )),
  inference(resolution,[],[f10,f13])).
tff(f13,plain,(
  attacked(germany)),
  inference(cnf_transformation,[],[f4])).
tff(f4,axiom,(
  attacked(germany)),
  file('/tmp/_tuzWjGqvL/SOT_E5Zjwa',unknown)).
tff(f10,plain,(
  ( ! [X0 : country,X1 : country] : (~attacked(X1) | protects(X0,X1) | ~member_of_nato(X1) | ~member_of_nato(X0)) )),
  inference(cnf_transformation,[],[f9])).
tff(f9,plain,(
  ! [X0 : country,X1 : country] : (protects(X0,X1) | ~attacked(X1) | ~member_of_nato(X1) | ~member_of_nato(X0))),
  inference(flattening,[],[f8])).
tff(f8,plain,(
  ! [X0 : country,X1 : country] : (protects(X0,X1) | (~attacked(X1) | ~member_of_nato(X1) | ~member_of_nato(X0)))),
  inference(ennf_transformation,[],[f1])).
tff(f1,axiom,(
  ! [X0 : country,X1 : country] : ((attacked(X1) & member_of_nato(X1) & member_of_nato(X0)) => protects(X0,X1))),
  file('/tmp/_tuzWjGqvL/SOT_E5Zjwa',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-RT/NonTrivial_03.THM.txt,True,0.95,1,True,0.95,1,THM,Unknown,1,0,0,Google---gemini-2.5-flash,Google---gemini-2.5-flash,OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% Converged at ZigZag number 1:1 with similarity 0.95
% The original and final NL have similarity: 0.95
% NL2L model:       Google---gemini-2.5-flash
% L2NL model:       Google---gemini-2.5-flash
% Similarity model: OpenAI---gpt-5-chat-latest
% -------------------------------------------------------------------------
% The original NL is :
If a country is a member of NATO, it will protect any member that is attacked.
Sweden is a member of NATO.
Germany is a member of NATO
Russia attacked Germany.
Will Sweden protect Germany?

% -------------------------------------------------------------------------
% The final NL is :
`country` is a type.
`member_of_NATO` is a predicate that takes a country.
`protects` is a predicate that takes two countries.
`attacked` is a predicate that takes two countries.
`sweden` is a country.
`germany` is a country.
`russia` is a country.
For any country C and any country M, if C is a member of NATO, and M is a member of NATO, and there exists a country A such that A attacked M, then C protects M.
Sweden is a member of NATO.
Germany is a member of NATO.
Russia attacked Germany.
Does Sweden protect Germany?
% -------------------------------------------------------------------------
% The differences between the original and final NL are:
Both segments express the same logical content: NATO members protect other NATO members that are attacked, and the question is whether Sweden will protect Germany after Russia attacks Germany.  
The first segment uses natural language, while the second formalizes the same information in logical or predicate notation. The only minor difference is that the second explicitly defines types and predicates, making the structure more formal but not changing the meaning.
% -------------------------------------------------------------------------
% The final logic is :
tff(country_type,type,
    country: $tType ).

tff(member_of_NATO_type,type,
    member_of_NATO: country > $o ).

tff(protects_type,type,
    protects: ( country * country ) > $o ).

tff(attacked_type,type,
    attacked: ( country * country ) > $o ).

tff(sweden_type,type,
    sweden: country ).

tff(germany_type,type,
    germany: country ).

tff(russia_type,type,
    russia: country ).

tff(axiom_country_protects_member,axiom,
    ! [C: country,M: country] :
      ( ( member_of_NATO(C)
        & member_of_NATO(M)
        & ? [A: country] : attacked(A,M) )
     => protects(C,M) ) ).

tff(axiom_sweden_is_member,axiom,
    member_of_NATO(sweden) ).

tff(axiom_germany_is_member,axiom,
    member_of_NATO(germany) ).

tff(axiom_russia_attacked_germany,axiom,
    attacked(russia,germany) ).

tff(conjecture_sweden_protects_germany,conjecture,
    protects(sweden,germany) ).

%
% -------------------------------------------------------------------------
% The ATP SZS status is THM with output
tff(type_def_5, type, country: $tType).
tff(func_def_0, type, sweden: country).
tff(func_def_1, type, germany: country).
tff(func_def_2, type, russia: country).
tff(pred_def_1, type, member_of_NATO: country > $o).
tff(pred_def_2, type, protects: (country * country) > $o).
tff(pred_def_3, type, attacked: (country * country) > $o).
tff(pred_def_4, type, sP0: country > $o).
tff(f21,plain,(
  $false),
  inference(finite_model_not_found_:_exhaustively_excluded_all_possible_domain_size_assignments,[],[f10,f11,f12,f13,f14])).
tff(f14,plain,(
  ~protects(sweden,germany)),
  inference(cnf_transformation,[],[f7])).
tff(f7,plain,(
  ~protects(sweden,germany)),
  inference(flattening,[],[f6])).
tff(f6,negated_conjecture,(
  ~protects(sweden,germany)),
  inference(negated_conjecture,[status(cth)],[f5])).
tff(f5,conjecture,(
  protects(sweden,germany)),
  file('/tmp/KkZmZRlY1x/SOT_8ZFDz4',unknown)).
tff(f13,plain,(
  attacked(russia,germany)),
  inference(cnf_transformation,[],[f4])).
tff(f4,axiom,(
  attacked(russia,germany)),
  file('/tmp/KkZmZRlY1x/SOT_8ZFDz4',unknown)).
tff(f12,plain,(
  member_of_NATO(germany)),
  inference(cnf_transformation,[],[f3])).
tff(f3,axiom,(
  member_of_NATO(germany)),
  file('/tmp/KkZmZRlY1x/SOT_8ZFDz4',unknown)).
tff(f11,plain,(
  member_of_NATO(sweden)),
  inference(cnf_transformation,[],[f2])).
tff(f2,axiom,(
  member_of_NATO(sweden)),
  file('/tmp/KkZmZRlY1x/SOT_8ZFDz4',unknown)).
tff(f10,plain,(
  ( ! [X2 : country,X0 : country,X1 : country] : (protects(X0,X1) | ~attacked(X2,X1) | ~member_of_NATO(X1) | ~member_of_NATO(X0)) )),
  inference(cnf_transformation,[],[f9])).
tff(f9,plain,(
  ! [X0 : country,X1 : country] : (protects(X0,X1) | ! [X2 : country] : ~attacked(X2,X1) | ~member_of_NATO(X1) | ~member_of_NATO(X0))),
  inference(flattening,[],[f8])).
tff(f8,plain,(
  ! [X0 : country,X1 : country] : (protects(X0,X1) | (! [X2 : country] : ~attacked(X2,X1) | ~member_of_NATO(X1) | ~member_of_NATO(X0)))),
  inference(ennf_transformation,[],[f1])).
tff(f1,axiom,(
  ! [X0 : country,X1 : country] : ((? [X2 : country] : attacked(X2,X1) & member_of_NATO(X1) & member_of_NATO(X0)) => protects(X0,X1))),
  file('/tmp/KkZmZRlY1x/SOT_8ZFDz4',unknown)).
% -------------------------------------------------------------------------
RESULT: TestFiles/SUMO-RT/NonTrivial_03.THM.txt,True,0.95,1,True,0.95,1,THM,Unknown,1,0,0,Google---gemini-2.5-flash,Google---gemini-2.5-flash,OpenAI---gpt-5-chat-latest
